```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

Output:
```
~/opt/anaconda3/lib/python3.9/site-packages/pandas/_testing/asserters.py in <module>
     11 )
     12 from pandas._libs.missing import is_matching_na
---> 13 import pandas._libs.testing as _testing
     14 
     15 from pandas.core.dtypes.common import (

~/opt/anaconda3/lib/python3.9/importlib/_bootstrap.py in parent(self)

KeyboardInterrupt:
```

# A1 | Week 3: Practice While Loops

------

### Instructions:

- Complete 50 points worth of questions to pass the assessment
  - You can attempt any number of questions and in any order provided you pass at least 50 points.
- Develop and check your answers in the spaces provided.
  - Do not remove any variables other provided markings already present in the answer spaces.
  - Do not make any changes to this notebook outside of the answer spaces provided.
  
#### Submitting

- Upload your .ipynb file to Gradescope (Make sure to restart the kernel and run all before you submit)
- You can submit as many times as needed
- When reviewing results, ignore any results listed under "Public Tests"

For more information, see the assessment page. We have also provided an example question below for clarity. 

#### Example Question
Write a Python program to display the first 10 integer numbers starting with 1 using while loops. Create the variable `example_output` and use it inside the while loop to construct the string that you want to display at the end of the program.

```
#This solution is for demonstration purposes ONLY. 
#There may be other ways of doing this task using while loops that produce the same output.
#You have the freedom to use a solution that you prefer as long as it uses while loop and 
#produces the expected output using the example_output variable.
example_output = ""
example_count = 1
while example_count <= 10:
    example_output += str(example_count)
    example_output += "\n"
    example_count += 1
example_output = example_output.strip() #this is to remove the extra new line in 10
print(example_output)
```

#### Question 1&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(1 Point)
Write a Python program to display integer numbers starting with 190 to 200 using while loops. Create the variable `output1` and use it inside the while loop to construct the string that you want to display at the end of the program.

Expected Output:   
190  
191  
192  
193  
194  
195  
196  
197  
198  
199  
200

```python
output1 = ""

# Write your solution here
temp = 190
line_break = "\n"
while(temp <= 200):
    output1 += str(temp) + line_break
    temp += 1

output1 = output1.strip()
print(output1)

```

#### Question 2&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(1 Point)
Write a Python program to display the first 10 even numbers starting with 0 using while loops. Create the variable `output2` and use  it inside the while loop to construct the string that you want to display at the end of the program.

*Note: The displayed string should be in the same format as in Q1. For clarity, the expected output is shown below.*

Expected Output:
<p>
0<br> 
2<br> 
4<br> 
6<br> 
8<br> 
10<br> 
12<br>
14<br> 
16<br> 
18
</p>

```python
output2 = ""

# Write your solution here
temp = 0
count = 0
line_break = "\n"
while(count < 10):
    output2 += str(temp) + line_break
    count += 1
    temp += 2

output2 = output2.strip()
print(output2)

```

#### Question 3&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(1 Point)
Write a Python program to display the following series using while loops: `30, 40, 50, 60, 70, 80, 90, ..., 300`. Create the variable `output3` and use  it inside the while loop to construct the string that you want to display at the end of the program.

*Note: The displayed string should be in the same format as in Q1. For clarity, the expected output is shown below.*

Expected Output:
<p>
30<br> 
40<br> 
50<br> 
60<br> 
70<br> 
80<br> 
90<br> 
...<br> 
300
</p>

```python
output3 = ""

# Write your solution here
temp = 30
line_break = "\n"
while(temp <= 300):
    output3 += str(temp) + line_break
    temp += 10

output3 = output3.strip()
print(output3)

```

#### Question 4&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(1 Point)
Write a Python program to display the squares of integer numbers divisible by 4 until 200 (inclusive) using while loops. Create the variable `output4` and use it inside the while loop to construct the string that you want to display at the end of the program. (Expected output series: 16, ..., 40000)

*Note: The displayed string should be in the same format as in Q1.*

```python
output4 = ""

# Write your solution here
temp = 1
line_break = "\n"

while temp <= 200:
    if (temp % 4 == 0):
        output4 += str(temp**2) + line_break
    temp += 1

output4 = output4.strip()
print(output4)

```

#### Question 5&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(1 Point)
Write a Python program to print the following series using while loops: `120, 116, 112, ..., 36`. Create the variable `output5` and use it inside the while loop to construct the string that you want to display at the end of the program.

*Note: The displayed string should be in the same format as in Q1.*

```python
output5 = ""

# Write your solution here
temp = 120
line_break = "\n"
while temp >= 36:
    output5 += str(temp) + line_break
    temp -= 4

output5 = output5.strip()
print(output5)

```

#### Question 6&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points)
Write a Python program that adds all the odd numbers until 300 (inclusive) using while loops. Create the variable `ouput6` and use it inside the while loop for the calculations needed. Print `output6` at the end of the program to check if it works correctly.

```python
output6 = 0

# Write your solution here
temp = 0
while (temp <= 300):
    if (temp % 2 != 0):
        output6 += temp
    temp += 1

print(output6)

```

#### Question 7&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points) 
Write a Python program to calculate the total of first 10 integer numbers using while loops. However, if the number is divisible by 5, do not use them to calculate the total value. Create the variable `output7` and use it inside the while loop for the calculations needed. [Expected output: 40]
***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
* You will need to use `continue` statement. This statement rejects all the remaining statements in the current iteration of the loop and moves the control back to the top of the loop.
</details>

```python
output7 = 0

# Write your solution here
temp = 1
while (temp <= 10):
    if (temp % 5 != 0):
        output7 += temp
    temp += 1

print(output7)

```

#### Question 8&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points) 
Write a Python program to calculate the total of first 20 integer numbers using while loops. However, if the program encounters a number that is divisible by 10 between 1 to 20, stop looping at the very moment and  calculate the total value until that point (i.e., excluding the number that is divisible by 10 and the numbers after that are also excluded). Create the variables `output8` and use it in the program for the calculations needed. [Expected output: 45]
***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
* You will need to use `break` statement. This will terminate the current loop.
</details>

```python
output8 = 0

# Write your solution here
temp = 1
while (temp <= 20):
    if (temp % 10 == 0):
        break
    output8 += temp
    temp += 1
print(output8)

```

#### Question 9&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points)
Write a Python program to display the average (rounded to two decimal places) of all the numbers which are divisible by 13 but not by 3 between 300 and 600.(exclusive both numbers) using while loop. Use the variable `output9` to assign the solution to it.

```python
# Write your solution here
temp = 301
sum = 0
total_numbers = 0;
while (temp < 600):
    if (temp % 13 == 0) and (temp % 3 != 0):
        sum += temp
        total_numbers += 1
    temp += 1
    
output9 = sum/total_numbers
output9 = round(output9, 2)
print(output9)

```

#### Question 10&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points)
Write a Python program to display the following pattern using while loops. Create the variable `output10` and use it inside the while loop to construct the string that you want to display at the end of the program.
<p>
*<br>
**<br>
***<br>
****<br>
*****<br>
</p>

```python
output10 = ""

# Write your solution here
row = 1
temp = "*"
line_break = "\n"
while (row <= 5):
    output10 += temp*row + line_break
    row += 1

output10 = output10.strip()
print(output10)

```

#### Question 11&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Write a Python program to display the following pattern using while loops. Create the variable `output11` and use it inside the while loop to construct the string that you want to display at the end of the program.
<p>
1<br>
22<br>
333<br>
4444<br>
55555<br>
666666<br>
7777777<br>
88888888<br>
999999999<br>
</p>

```python
output11 = ""

# Write your solution here
row = 1
line_break = "\n"

while (row <= 9):
    output11 += str(row)*row + line_break
    row += 1
    
output11 = output11.strip()
print(output11)

```

#### Question 12&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Write a Python program to display numbers that are divisible by 9 and multiple of 5, between 1000 and 2000 (including 1000 and 2000) using while loops. Create the variable `output12` and use it inside the while loop to construct the string that you want to display at the end of the program.

*Note: The displayed string should be in the same format as in Q1.*

```python
output12 = ""

# Write your solution here
temp = 1000
line_break = "\n"

while (temp <= 2000):
    if (temp % 9 == 0) and (temp % 5 == 0):
        output12 += str(temp) + line_break
    temp += 1
    
output12 = output12.strip()
print(output12)

```

#### Question 13&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Write a Python program that iterates from 200 to 250 (including both 200 and 250) using while loops and perform the following. Create the variable `output13` and use it inside the while loop to construct the string that you want to display at the end of the program.

1. If a number is a multiply of 3, display "NUMBER is a multiply of 3!" (*Note: NUMBER should be replaced with the value at that iteration*)
2. If a number is a multiply of 5, display "NUMBER is a multiply of 5!" (*Note: NUMBER should be replaced with the value at that iteration*)
3. If a number is a multiply of 3 as well as multiply of 5, display "NUMBER is a multiply of both 3 and 5!" (*Note: NUMBER should be replaced with the value at that iteration*)

*Note: The displayed string should be in the same format as in Q1.*

```python
output13 = ""

# Write your solution here
temp = 200
while (temp <= 250):
    if (temp % 3 == 0) and (temp % 5 == 0):
        output13 += f"{temp} is a multiply of both 3 and 5!\n"
        
    elif (temp % 3 == 0):
        output13 += f"{temp} is a multiply of 3!\n"
        
    elif (temp % 5 == 0):
        output13 += f"{temp} is a multiply of 5!\n"
        
    temp += 1

output13 = output13.strip()
print(output13)
```

#### Question 14&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Create the variable `q14_num` and assign 1912 to it. Write a Python program to reverse `q14_num` using while loop and arithmetic operations. For example, if `q14_num` is 1912, the reversed number is 2191. Create the variable `output14` and assign the reversed number to it. However, if a user has mistakenly assign one digit number (i.e., 0-9) to `q14_num`, `output14`should be the string "Please enter at least two digits to perform operations". Print `output14` at the end of the prohram to check if it works correctly.

```python
q14_num = 1912

# Write your solution here
if (q14_num >= 0 and q14_num <= 9):
    output14 = "Please enter at least two digits to perform operations"
else:
    output14 = int(str(q14_num)[::-1])

print(output14)

```

#### Question 15 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points) 
Write a Python program to display the series of `2, 22, 222, 2222, ..., n terms` till n terms using while loops.

For this question, create `q15_n` variable (that represents `n` in the series) and assign 20 to it. Create the variable `output15` and use it inside the while loop to construct the string that you want to display at the end of the program.

*Note: The displayed string should be in the same format as in Q1.*

```python
q15_n = 20
output15 = ""

# Write your solution here
temp = 1
line_break = "\n"

while (temp <= q15_n):
    output15 += "2"*temp + line_break
    temp += 1

output15 = output15.strip()
print(output15)

```

#### Question 16 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Write a Python program to find the sum of following series: 1 + 4 - 9 + 16 - 25 + 36 ... n terms.

For this question, create `q16_n` variable (that represents `n` in the series) and assign 20 to it. Create the variable `output16` and assign the solution to it.

```python
q16_n = 20

# Write your solution here
temp = 1
output16 = 0
while (temp <= q16_n):
    
    if (temp == 1):
         output16 += temp**2
            
    elif (temp % 2 == 0):
        output16 += temp**2
        
    else:
        output16 -= temp**2
        
    temp += 1

print(output16)
```

Output:
```
212
```

#### Question 17 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
 Write a Python program to calculate factorial using while loop. Create a variable `q17_input` and assign the number 191 to it. To calculate the factorial use the following guidelines. Create a variable `output17` and assign the solution to it.

* Factorial of number n is represented by `n!` and it is equal to `1*2*3*....*(n-1)*n`. For example, `4!` = `1*2*3*4` = 24
* Note that 1! = 1 and 0! = 1

```python
q17_input = 191

# Write your solution here
output17 = 1
if (q17_input == 0):
    output17 = 0
else:
    temp = 1
    while (temp <= q17_input):
        output17 *= temp
        temp += 1
        
print(output17)

```

#### Question 18 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python program to display the following pattern using while loops. Create the variable `output18` and use it inside the while loop to construct the string that you want to display at the end of the program.
<p>
    &nbsp;&nbsp;&nbsp;&nbsp;*  <br>
  &nbsp;&nbsp;&nbsp;*** <br>
  &nbsp;&nbsp;***** <br>
&nbsp;******* <br>
********* <br>
&nbsp;******* <br>
&nbsp;&nbsp;***** <br>
   &nbsp;&nbsp;&nbsp;*** <br>
    &nbsp;&nbsp;&nbsp;&nbsp;*  <br>
</p>

***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
* You may need to use two while loops to print the first half and second half of the pattern.
* You may need to use `break` statement as practiced in Question 8.
</details>

```python
output18 = ""

# Write your solution here

import math

# init value
row = 1
number_of_wildcards = 1
space = " "
wildcard = "*"

# Change max number of wildcards
max_wildcards = 9
half_rows = math.ceil(max_wildcards/2)

while (row <= max_wildcards):
    
    if (row <= half_rows):
        # First half
        while (number_of_wildcards % 2 == 0):
            number_of_wildcards += 1

        number_of_spaces_beginning_ending = int((max_wildcards - number_of_wildcards) / 2)
        output18 += space * number_of_spaces_beginning_ending +  wildcard * number_of_wildcards + space * number_of_spaces_beginning_ending + "\n"
        number_of_wildcards += 1
        
    else:
        # Second half
        while (number_of_wildcards % 2 == 0 or number_of_wildcards == max_wildcards):
            number_of_wildcards -= 1

        number_of_spaces_beginning_ending = int((max_wildcards - number_of_wildcards) / 2)
        output18 += space * number_of_spaces_beginning_ending +  wildcard * number_of_wildcards + space * number_of_spaces_beginning_ending + "\n"
        number_of_wildcards -= 1
    
    row += 1

output18 = output18.rstrip()
print(output18)
```

#### Question 19 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python program to reverse a given number using while loop. Create a variable `q19_input` and assign the number 121019200109 to it. Use the variable `output19` and assign the solution to it.

```python
q19_input = 121019200109

# Write your solution here
output19 = ""
while (q19_input > 0):
    temp = q19_input % 10
    output19 += str(temp)
    q19_input = q19_input // 10

output19 = int(output19)

print(output19)
```

#### Question 20 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python program to find the sum of the following series using a while loop.  
$1 +\frac{y}{1!} + \frac{y^2}{2!} + ... + \frac{y^n}{n!}$

Use the following guidelines when developing the program.

* Create two variables `q20_y` and `q20_n` and assign the values 6 and 26 (Note that `q20_y` and `q20_n` represent `y` and `n` in the above-mentioned series).
* The sum should be ronded to the nearest integer value. Create the variable `output20` and assign the solution to it.

```python
q20_y, q20_n = 6, 26

# Write your solution here
temp_n = 1
output20 = 1

while (temp_n <= q20_n):
    # Calculate factorial number of temp_n
    index = 1
    fractorial_of_temp_n = 1
    
    while (index <= temp_n):
        fractorial_of_temp_n *= index
        index += 1
    
    # Calculate output20
    output20 += (q20_y**temp_n)/fractorial_of_temp_n
    
    # Increase temp_n by 1
    temp_n += 1

# The sum should be ronded to the nearest integer value
output20 = int(round(output20, 0))
print(output20)

```

#### Question 21 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python program to check whether a number is palindrome or not using while loop.
Palindrome is a sequence which reads the same backward as forward. Examples of palindromes are mentioned below.
* 12344321
* 2002
* 19191

Create the variable `q21_input` and assign 12123454321 to it. If `q21_input` is a palindrome assign the string "This is a palindrome" to `output21`. Otherwise, `output21` should be "This is not a palindrome"

```python
q21_input = 12123454321

# Write your solution here
revert_of_q21_input = int(str(q21_input)[::-1])
if (q21_input == revert_of_q21_input):
    output21 = "This is a palindrome"
else:
    output21 = "This is not a palindrome"
    
print(output21)

```

#### Question 22 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python program that displays sum of odd numbers and even numbers separately that fall between two given numbers `q22_num1` and `q22_num2` (including both numbers) using while loop. Use the following guidelines to develop this program.

* For this question, assign the values 21 and 9 to the two variables `q22_num1` and `q22_num2`, respectively.
* Use the two variables `output22_odd` and `output22_even` to assign the sum of odd numbers and the sum of even numbers, respectively.

```python
q22_num1 = 21
q22_num2 = 9

# Write your solution here
q22_max = max(q22_num1, q22_num2)
q22_min = min(q22_num1, q22_num2)

output22_odd = 0
output22_even = 0
temp = q22_min
while (temp <= q22_max):
    if (temp % 2 == 0):
        output22_even += temp
    else:
        output22_odd += temp
    temp += 1

print(output22_even)
print(output22_odd)

```

#### Question 23 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(10 Points)
Write a Python program to display the following pattern using while loops. Create the variable `output23` and use it inside the while loop to construct the string that you want to display at the end of the program.
<p>
 &nbsp;***<br>                                                                 
 *&nbsp;&nbsp;&nbsp;*<br>                                                                  
 *&nbsp;&nbsp;&nbsp;*<br>                                                                  
 *&nbsp;&nbsp;&nbsp;*<br>                                                                  
 *&nbsp;&nbsp;&nbsp;*<br>                                                                  
 *&nbsp;&nbsp;&nbsp;*<br>                                                                  
&nbsp;*** 
</p>

```python
output23 = "" 

# Write your solution here
row = 7
temp = 1
while (temp <= row):
    if (temp == 1):
        output23 += "  ***  " + "\n"
    elif (temp == row):
        output23 += "  ***"
    else:
        output23 += " *   * " + "\n"
    temp += 1

print(output23)
```

#### Question 24 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(10 Points)
Write a Python program to display the following pattern using while loops. Create the variable `output24` and use it inside the while loop to construct the string that you want to display at the end of the program.

<p>
    * * * * *<br> 
&nbsp;* * * *<br> 
&nbsp;&nbsp;* * *<br> 
&nbsp;&nbsp;&nbsp;* *<br> 
&nbsp;&nbsp;&nbsp;&nbsp;*<br> 
&nbsp;&nbsp;&nbsp;&nbsp;*<br> 
&nbsp;&nbsp;&nbsp;* *<br> 
&nbsp;&nbsp;* * *<br> 
&nbsp;* * * *<br> 
* * * * *<br>
</p>

```python
output24 = ""

# Write your solution here

import math

# init value
space = " "
wildcard = "*"
sample_line = ""

# Change max number of wildcards
max_wildcards = 5

# generate sample line (line with max wildcards)
temp = 1
while (temp <= max_wildcards):
    sample_line += wildcard + space
    temp += 1
max_length = len(sample_line)

# generate first half of output24
number_of_wildcards = 0
temp_line = sample_line
# First half
while (number_of_wildcards < max_wildcards):
    if (number_of_wildcards != 0):
        # Calculate new sample line
        temp_line = temp_line[number_of_wildcards + 1:len(temp_line)]
        
        # Add space into sample line
        temp_line = space*number_of_wildcards + temp_line
        
    output24 += temp_line + "\n"
    number_of_wildcards += 1

# Append reverse of output24
splits = output24.split("\n")
index = len(splits) - 1

while (index >= 0):
    if (splits[index] != ""):
        output24 += splits[index] + "\n"
    index -= 1

output24 = output24.rstrip()
print(output24)
```

Tests for Question 01 (points: 1)


Tests for Question 02 (points: 1)


Tests for Question 03 (points: 1)


Tests for Question 04 (points: 1)


Tests for Question 05 (points: 1)


Tests for Question 06 (points: 3)


Tests for Question 07 (points: 3)


Tests for Question 08 (points: 3)


Tests for Question 09 (points: 3)


Tests for Question 10 (points: 3)


Tests for Question 11 (points: 4)


Tests for Question 12 (points: 4)


Tests for Question 13 (points: 4)


Tests for Question 14 (points: 4)


Tests for Question 15 (points: 4)


Tests for Question 16 (points: 4)


Tests for Question 17 (points: 6)


Tests for Question 18 (points: 6)


Tests for Question 19 (points: 6)


Tests for Question 20 (points: 6)


Tests for Question 21 (points: 6)


Tests for Question 22 (points: 6)
```
>>> output22_even
90
```


Tests for Question 23 (points: 10)


Tests for Question 24 (points: 10)
